<!DOCTYPE html PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN" "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml">
<head>
<meta http-equiv="Content-Type" content="text/xhtml;charset=UTF-8"/>
<meta http-equiv="X-UA-Compatible" content="IE=9"/>
<title>Himalaya Service Master: /home/ldx/workspaces/Himalaya/ServiceSW/Master/Components/SystemTracking/Source/MainControl.cpp Source File</title>

<link href="tabs.css" rel="stylesheet" type="text/css"/>
<link href="doxygen.css" rel="stylesheet" type="text/css" />
<link href="navtree.css" rel="stylesheet" type="text/css"/>
<script type="text/javascript" src="jquery.js"></script>
<script type="text/javascript" src="resize.js"></script>
<script type="text/javascript" src="navtree.js"></script>
<script type="text/javascript">
  $(document).ready(initResizable);
</script>
<link href="search/search.css" rel="stylesheet" type="text/css"/>
<script type="text/javascript" src="search/search.js"></script>
<script type="text/javascript">
  $(document).ready(function() { searchBox.OnSelectItem(0); });
</script>

</head>
<body>
<div id="top"><!-- do not remove this div! -->


<div id="titlearea">
<table cellspacing="0" cellpadding="0">
 <tbody>
 <tr style="height: 56px;">
  
  
  <td style="padding-left: 0.5em;">
   <div id="projectname">Himalaya Service Master
   
   </div>
   
  </td>
  
  
  
 </tr>
 </tbody>
</table>
</div>

<!-- Generated by Doxygen 1.7.6.1 -->
<script type="text/javascript">
var searchBox = new SearchBox("searchBox", "search",false,'Search');
</script>
  <div id="navrow1" class="tabs">
    <ul class="tablist">
      <li><a href="index.html"><span>Main&#160;Page</span></a></li>
      <li><a href="annotated.html"><span>Classes</span></a></li>
      <li class="current"><a href="files.html"><span>Files</span></a></li>
      <li><a href="dirs.html"><span>Directories</span></a></li>
      <li>
        <div id="MSearchBox" class="MSearchBoxInactive">
        <span class="left">
          <img id="MSearchSelect" src="search/mag_sel.png"
               onmouseover="return searchBox.OnSearchSelectShow()"
               onmouseout="return searchBox.OnSearchSelectHide()"
               alt=""/>
          <input type="text" id="MSearchField" value="Search" accesskey="S"
               onfocus="searchBox.OnSearchFieldFocus(true)" 
               onblur="searchBox.OnSearchFieldFocus(false)" 
               onkeyup="searchBox.OnSearchFieldChange(event)"/>
          </span><span class="right">
            <a id="MSearchClose" href="javascript:searchBox.CloseResultsWindow()"><img id="MSearchCloseImg" border="0" src="search/close.png" alt=""/></a>
          </span>
        </div>
      </li>
    </ul>
  </div>
  <div id="navrow2" class="tabs2">
    <ul class="tablist">
      <li><a href="files.html"><span>File&#160;List</span></a></li>
    </ul>
  </div>
</div>
<div id="side-nav" class="ui-resizable side-nav-resizable">
  <div id="nav-tree">
    <div id="nav-tree-contents">
    </div>
  </div>
  <div id="splitbar" style="-moz-user-select:none;" 
       class="ui-resizable-handle">
  </div>
</div>
<script type="text/javascript">
  initNavTree('MainControl_8cpp.html','');
</script>
<div id="doc-content">
<div class="header">
  <div class="headertitle">
<div class="title">MainControl.cpp</div>  </div>
</div><!--header-->
<div class="contents">
<a href="MainControl_8cpp.html">Go to the documentation of this file.</a><div class="fragment"><pre class="fragment"><a name="l00001"></a>00001 <span class="comment">/****************************************************************************/</span><span class="comment"></span>
<a name="l00002"></a>00002 <span class="comment">/*! \file MainControl.cpp</span>
<a name="l00003"></a>00003 <span class="comment"> *</span>
<a name="l00004"></a>00004 <span class="comment"> *  \brief Implementation of Main control configuration class.</span>
<a name="l00005"></a>00005 <span class="comment"> *</span>
<a name="l00006"></a>00006 <span class="comment"> *   $Version: $ 0.1</span>
<a name="l00007"></a>00007 <span class="comment"> *   $Date:    $ 2013-05-28</span>
<a name="l00008"></a>00008 <span class="comment"> *   $Author:  $ R.Wu</span>
<a name="l00009"></a>00009 <span class="comment"> *</span>
<a name="l00010"></a>00010 <span class="comment"> *  \b Company:</span>
<a name="l00011"></a>00011 <span class="comment"> *</span>
<a name="l00012"></a>00012 <span class="comment"> *       Leica Biosystems R&amp;D Center Shanghai.</span>
<a name="l00013"></a>00013 <span class="comment"> *</span>
<a name="l00014"></a>00014 <span class="comment"> *  (C) Copyright 2010 by LBS R&amp;D Center Shanghai. All rights reserved.</span>
<a name="l00015"></a>00015 <span class="comment"> *  This is unpublished proprietary source code of Leica. The copyright notice</span>
<a name="l00016"></a>00016 <span class="comment"> *  does not evidence any actual or intended publication.</span>
<a name="l00017"></a>00017 <span class="comment"> *</span>
<a name="l00018"></a>00018 <span class="comment"> */</span>
<a name="l00019"></a>00019 <span class="comment">/****************************************************************************/</span>
<a name="l00020"></a>00020 
<a name="l00021"></a>00021 <span class="preprocessor">#include &quot;<a class="code" href="MainControl_8h.html" title="Declaration of Main control configuration class.">SystemTracking/Include/MainControl.h</a>&quot;</span>
<a name="l00022"></a>00022 
<a name="l00023"></a>00023 <span class="preprocessor">#include &lt;QDebug&gt;</span>
<a name="l00024"></a>00024 
<a name="l00025"></a>00025 <span class="preprocessor">#include &quot;<a class="code" href="DlgModifySubModule_8h.html" title="Declaration of Modify module dialog.">SystemTracking/Include/DlgModifySubModule.h</a>&quot;</span>
<a name="l00026"></a>00026 <span class="preprocessor">#include &quot;Main/Include/HimalayaServiceEventCodes.h&quot;</span>
<a name="l00027"></a>00027 <span class="preprocessor">#include &quot;Global/Include/Utils.h&quot;</span>
<a name="l00028"></a>00028 
<a name="l00029"></a>00029 <span class="preprocessor">#include &quot;ui_MainControl.h&quot;</span>
<a name="l00030"></a>00030 
<a name="l00031"></a>00031 <span class="preprocessor">#include &quot;ServiceDataManager/Include/TestCaseGuide.h&quot;</span>
<a name="l00032"></a>00032 <span class="preprocessor">#include &quot;ServiceDataManager/Include/TestCase.h&quot;</span>
<a name="l00033"></a>00033 <span class="preprocessor">#include &quot;ServiceDataManager/Include/TestCaseFactory.h&quot;</span>
<a name="l00034"></a>00034 
<a name="l00035"></a>00035 <span class="keyword">namespace </span>SystemTracking {
<a name="l00036"></a>00036 
<a name="l00037"></a>00037 <span class="keyword">const</span> QString MODULE_MAINCONTROL(<span class="stringliteral">&quot;Main Control&quot;</span>);    <span class="comment">//!&lt; Main Control name</span>
<a name="l00038"></a>00038 <span class="comment"></span><span class="keyword">const</span> QString SUBMODULE_EBOX(<span class="stringliteral">&quot;E Box&quot;</span>);               <span class="comment">//!&lt; name of E box submodule</span>
<a name="l00039"></a>00039 <span class="comment"></span><span class="keyword">const</span> QString SUBMODULE_ASB3(<span class="stringliteral">&quot;ASB3&quot;</span>);                <span class="comment">//!&lt; name of ASB3 submodule</span>
<a name="l00040"></a>00040 <span class="comment"></span><span class="keyword">const</span> QString SUBMODULE_ASB5(<span class="stringliteral">&quot;ASB5&quot;</span>);                <span class="comment">//!&lt; name of ASB5 submodule</span>
<a name="l00041"></a>00041 <span class="comment"></span><span class="keyword">const</span> QString SUBMODULE_ASB15(<span class="stringliteral">&quot;ASB15&quot;</span>);              <span class="comment">//!&lt; name of ASB15 submodule</span>
<a name="l00042"></a>00042 <span class="comment"></span><span class="keyword">const</span> QString SUBMODULE_TOUCHSCREEN(<span class="stringliteral">&quot;Touch Screen&quot;</span>); <span class="comment">//!&lt; name of touch screen submodule</span>
<a name="l00043"></a>00043 <span class="comment"></span><span class="keyword">const</span> QString SUBMODULE_VENTFAN(<span class="stringliteral">&quot;Ventilation Fan&quot;</span>);  <span class="comment">//!&lt; name of ventilation fan submodule</span>
<a name="l00044"></a>00044 <span class="comment"></span>
<a name="l00045"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#a71967298fb2e9f843c8e58051883e6e4">00045</a> <a class="code" href="classSystemTracking_1_1CMainControl.html#a71967298fb2e9f843c8e58051883e6e4" title="Default constructor.">CMainControl::CMainControl</a>(Core::CServiceGUIConnector &amp;DataConnector,
<a name="l00046"></a>00046                            QWidget *p_Parent)
<a name="l00047"></a>00047     : QWidget(p_Parent)
<a name="l00048"></a>00048     , mp_DateConnector(&amp;DataConnector)
<a name="l00049"></a>00049     , mp_ModuleList(NULL)
<a name="l00050"></a>00050     , mp_Ui(new Ui::<a class="code" href="classSystemTracking_1_1CMainControl.html" title="This is the widget displaying to modify main control config.">CMainControl</a>)
<a name="l00051"></a>00051 {
<a name="l00052"></a>00052     mp_Ui-&gt;setupUi(<span class="keyword">this</span>);
<a name="l00053"></a>00053 
<a name="l00054"></a>00054     mp_Ui-&gt;finalizeConfigBtn-&gt;setEnabled(<span class="keyword">false</span>);
<a name="l00055"></a>00055 
<a name="l00056"></a>00056     mp_MessageDlg = <span class="keyword">new</span> MainMenu::CMessageDlg(<span class="keyword">this</span>);
<a name="l00057"></a>00057     mp_MessageDlg-&gt;SetTitle(QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00058"></a>00058                                                     <span class="stringliteral">&quot;Finalize Configuration&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00059"></a>00059     mp_MessageDlg-&gt;setModal(<span class="keyword">true</span>);
<a name="l00060"></a>00060 
<a name="l00061"></a>00061     (void)connect(mp_Ui-&gt;modifyEBox,
<a name="l00062"></a>00062                   SIGNAL(clicked()),
<a name="l00063"></a>00063                   <span class="keyword">this</span>,
<a name="l00064"></a>00064                   SLOT(<a class="code" href="classSystemTracking_1_1CMainControl.html#a424c87c626dd641ff3df7fc3c9aa6a4e" title="slot for modify E box.">ModifyEBox</a>()) );
<a name="l00065"></a>00065 
<a name="l00066"></a>00066     (void)connect(mp_Ui-&gt;modifyASB3,
<a name="l00067"></a>00067                   SIGNAL(clicked()),
<a name="l00068"></a>00068                   <span class="keyword">this</span>,
<a name="l00069"></a>00069                   SLOT(<a class="code" href="classSystemTracking_1_1CMainControl.html#a4af4a3c0a79d59b6fc3edb7952928a1f" title="slot for modify ASB3.">ModifyASB3</a>()) );
<a name="l00070"></a>00070 
<a name="l00071"></a>00071     (void)connect(mp_Ui-&gt;modifyASB5,
<a name="l00072"></a>00072                   SIGNAL(clicked()),
<a name="l00073"></a>00073                   <span class="keyword">this</span>,
<a name="l00074"></a>00074                   SLOT(<a class="code" href="classSystemTracking_1_1CMainControl.html#ade32f77c4473db0eef5998fe1ab83fe9" title="slot for modify ASB5.">ModifyASB5</a>()) );
<a name="l00075"></a>00075 
<a name="l00076"></a>00076     (void)connect(mp_Ui-&gt;modifyASB15,
<a name="l00077"></a>00077                   SIGNAL(clicked()),
<a name="l00078"></a>00078                   <span class="keyword">this</span>,
<a name="l00079"></a>00079                   SLOT(<a class="code" href="classSystemTracking_1_1CMainControl.html#abed872e022218b735429921ff01e975e" title="slot for modify ASB15.">ModifyASB15</a>()) );
<a name="l00080"></a>00080 
<a name="l00081"></a>00081     (void)connect(mp_Ui-&gt;modifyTouchScreen,
<a name="l00082"></a>00082                   SIGNAL(clicked()),
<a name="l00083"></a>00083                   <span class="keyword">this</span>,
<a name="l00084"></a>00084                   SLOT(<a class="code" href="classSystemTracking_1_1CMainControl.html#aee506e100f3fff5b3a83b2090de9b859" title="slot for modify Touch screen.">ModifyTouchScreen</a>()) );
<a name="l00085"></a>00085 
<a name="l00086"></a>00086     (void)connect(mp_Ui-&gt;modifyVentFan,
<a name="l00087"></a>00087                   SIGNAL(clicked()),
<a name="l00088"></a>00088                   <span class="keyword">this</span>,
<a name="l00089"></a>00089                   SLOT(<a class="code" href="classSystemTracking_1_1CMainControl.html#ac56412bf8e60d469b8b32b563c33d437" title="slot for modify Vent fan.">ModifyVentFan</a>()) );
<a name="l00090"></a>00090 
<a name="l00091"></a>00091     (void)connect(mp_Ui-&gt;finalizeConfigBtn,
<a name="l00092"></a>00092                   SIGNAL(clicked()),
<a name="l00093"></a>00093                   <span class="keyword">this</span>,
<a name="l00094"></a>00094                   SLOT(<a class="code" href="classSystemTracking_1_1CMainControl.html#a341ffbc353fc0bac77318eae146bf4d1" title="slot for finalize configuration.">OnFinalizeConfiguration</a>()) );
<a name="l00095"></a>00095 }
<a name="l00096"></a>00096 
<a name="l00097"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#ae5f43a8c091afb7f0280a5e7ef493e5e">00097</a> <a class="code" href="classSystemTracking_1_1CMainControl.html#ae5f43a8c091afb7f0280a5e7ef493e5e" title="Destructor.">CMainControl::~CMainControl</a>()
<a name="l00098"></a>00098 {
<a name="l00099"></a>00099     <span class="keywordflow">try</span>
<a name="l00100"></a>00100     {
<a name="l00101"></a>00101         <span class="keyword">delete</span> mp_Ui;
<a name="l00102"></a>00102         <span class="keyword">delete</span> mp_MessageDlg;
<a name="l00103"></a>00103     }
<a name="l00104"></a>00104     <span class="keywordflow">catch</span> (...) { }
<a name="l00105"></a>00105 }
<a name="l00106"></a>00106 
<a name="l00107"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#a49a9bf2719340b4d072118f4f6925c33">00107</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#a49a9bf2719340b4d072118f4f6925c33" title="slot for update sub module.">CMainControl::UpdateSubModule</a>(ServiceDataManager::CSubModule &amp;SubModule)
<a name="l00108"></a>00108 {
<a name="l00109"></a>00109     <span class="keywordflow">if</span> (!mp_ModuleList) {
<a name="l00110"></a>00110         mp_ModuleList = <span class="keyword">new</span> ServiceDataManager::CModuleDataList;
<a name="l00111"></a>00111         ServiceDataManager::CModuleDataList* ModuleList = mp_DateConnector-&gt;GetModuleListContainer();
<a name="l00112"></a>00112         <span class="keywordflow">if</span> (!ModuleList) {
<a name="l00113"></a>00113             qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::UpdateSubModule(): Invalid module list!&quot;</span>;
<a name="l00114"></a>00114             <span class="keywordflow">return</span>;
<a name="l00115"></a>00115         }
<a name="l00116"></a>00116 
<a name="l00117"></a>00117         *mp_ModuleList = *ModuleList;
<a name="l00118"></a>00118     }
<a name="l00119"></a>00119 
<a name="l00120"></a>00120     ServiceDataManager::CModule *pModule = mp_ModuleList-&gt;GetModule(MODULE_MAINCONTROL);
<a name="l00121"></a>00121     <span class="keywordflow">if</span> (0 == pModule)
<a name="l00122"></a>00122     {
<a name="l00123"></a>00123         qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::UpdateSubModule(): Invalid module : &quot;</span>
<a name="l00124"></a>00124                  &lt;&lt; MODULE_MAINCONTROL;
<a name="l00125"></a>00125         <span class="keywordflow">return</span>;
<a name="l00126"></a>00126     }
<a name="l00127"></a>00127 
<a name="l00128"></a>00128     (void)pModule-&gt;UpdateSubModule(&amp;SubModule);
<a name="l00129"></a>00129 
<a name="l00130"></a>00130     mp_Ui-&gt;finalizeConfigBtn-&gt;setEnabled(<span class="keyword">true</span>);
<a name="l00131"></a>00131 }
<a name="l00132"></a>00132 
<a name="l00133"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#a424c87c626dd641ff3df7fc3c9aa6a4e">00133</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#a424c87c626dd641ff3df7fc3c9aa6a4e" title="slot for modify E box.">CMainControl::ModifyEBox</a>()
<a name="l00134"></a>00134 {
<a name="l00135"></a>00135     Global::EventObject::Instance().RaiseEvent(EVENT_GUI_MODIFY_MAIN_CONTROL_SUBMODULE,
<a name="l00136"></a>00136                                                Global::tTranslatableStringList()&lt;&lt;<span class="stringliteral">&quot;EBox&quot;</span>);
<a name="l00137"></a>00137     qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::ModifyEBox !&quot;</span>;
<a name="l00138"></a>00138 
<a name="l00139"></a>00139     this-&gt;ModifySubModule(MODULE_MAINCONTROL, SUBMODULE_EBOX);
<a name="l00140"></a>00140 }
<a name="l00141"></a>00141 
<a name="l00142"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#a4af4a3c0a79d59b6fc3edb7952928a1f">00142</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#a4af4a3c0a79d59b6fc3edb7952928a1f" title="slot for modify ASB3.">CMainControl::ModifyASB3</a>(<span class="keywordtype">void</span>)
<a name="l00143"></a>00143 {
<a name="l00144"></a>00144     Global::EventObject::Instance().RaiseEvent(EVENT_GUI_MODIFY_MAIN_CONTROL_SUBMODULE,
<a name="l00145"></a>00145                                                Global::tTranslatableStringList()&lt;&lt;<span class="stringliteral">&quot;ASB3&quot;</span>);
<a name="l00146"></a>00146     qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::ModifyASB3 !&quot;</span>;
<a name="l00147"></a>00147 
<a name="l00148"></a>00148     this-&gt;ModifySubModule(MODULE_MAINCONTROL, SUBMODULE_ASB3);
<a name="l00149"></a>00149 }
<a name="l00150"></a>00150 
<a name="l00151"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#ade32f77c4473db0eef5998fe1ab83fe9">00151</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#ade32f77c4473db0eef5998fe1ab83fe9" title="slot for modify ASB5.">CMainControl::ModifyASB5</a>(<span class="keywordtype">void</span>)
<a name="l00152"></a>00152 {
<a name="l00153"></a>00153     Global::EventObject::Instance().RaiseEvent(EVENT_GUI_MODIFY_MAIN_CONTROL_SUBMODULE,
<a name="l00154"></a>00154                                                Global::tTranslatableStringList()&lt;&lt;<span class="stringliteral">&quot;ASB5&quot;</span>);
<a name="l00155"></a>00155     qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::ModifyASB5 !&quot;</span>;
<a name="l00156"></a>00156 
<a name="l00157"></a>00157     this-&gt;ModifySubModule(MODULE_MAINCONTROL, SUBMODULE_ASB5);
<a name="l00158"></a>00158 }
<a name="l00159"></a>00159 
<a name="l00160"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#abed872e022218b735429921ff01e975e">00160</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#abed872e022218b735429921ff01e975e" title="slot for modify ASB15.">CMainControl::ModifyASB15</a>(<span class="keywordtype">void</span>)
<a name="l00161"></a>00161 {
<a name="l00162"></a>00162     Global::EventObject::Instance().RaiseEvent(EVENT_GUI_MODIFY_MAIN_CONTROL_SUBMODULE,
<a name="l00163"></a>00163                                                Global::tTranslatableStringList()&lt;&lt;<span class="stringliteral">&quot;ASB15&quot;</span>);
<a name="l00164"></a>00164     qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::ModifyASB15 !&quot;</span>;
<a name="l00165"></a>00165 
<a name="l00166"></a>00166     this-&gt;ModifySubModule(MODULE_MAINCONTROL, SUBMODULE_ASB15);
<a name="l00167"></a>00167 }
<a name="l00168"></a>00168 
<a name="l00169"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#aee506e100f3fff5b3a83b2090de9b859">00169</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#aee506e100f3fff5b3a83b2090de9b859" title="slot for modify Touch screen.">CMainControl::ModifyTouchScreen</a>(<span class="keywordtype">void</span>)
<a name="l00170"></a>00170 {
<a name="l00171"></a>00171     Global::EventObject::Instance().RaiseEvent(EVENT_GUI_MODIFY_MAIN_CONTROL_SUBMODULE,
<a name="l00172"></a>00172                                                Global::tTranslatableStringList()&lt;&lt;<span class="stringliteral">&quot;Touch Screen&quot;</span>);
<a name="l00173"></a>00173     qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::ModifyTouchScreen !&quot;</span>;
<a name="l00174"></a>00174 
<a name="l00175"></a>00175     this-&gt;ModifySubModule(MODULE_MAINCONTROL, SUBMODULE_TOUCHSCREEN);
<a name="l00176"></a>00176 }
<a name="l00177"></a>00177 
<a name="l00178"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#ac56412bf8e60d469b8b32b563c33d437">00178</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#ac56412bf8e60d469b8b32b563c33d437" title="slot for modify Vent fan.">CMainControl::ModifyVentFan</a>(<span class="keywordtype">void</span>)
<a name="l00179"></a>00179 {
<a name="l00180"></a>00180     Global::EventObject::Instance().RaiseEvent(EVENT_GUI_MODIFY_MAIN_CONTROL_SUBMODULE,
<a name="l00181"></a>00181                                                Global::tTranslatableStringList()&lt;&lt;<span class="stringliteral">&quot;Ventilation Fan&quot;</span>);
<a name="l00182"></a>00182 
<a name="l00183"></a>00183     this-&gt;ModifySubModule(MODULE_MAINCONTROL, SUBMODULE_VENTFAN);
<a name="l00184"></a>00184     qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::ModifyVentFan !&quot;</span>;
<a name="l00185"></a>00185 }
<a name="l00186"></a>00186 
<a name="l00187"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#a341ffbc353fc0bac77318eae146bf4d1">00187</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#a341ffbc353fc0bac77318eae146bf4d1" title="slot for finalize configuration.">CMainControl::OnFinalizeConfiguration</a>(<span class="keywordtype">void</span>)
<a name="l00188"></a>00188 {
<a name="l00189"></a>00189     <span class="comment">//save config to file</span>
<a name="l00190"></a>00190     QString Text = QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00191"></a>00191                                            <span class="stringliteral">&quot;Do you want to finalize the configuration for the Main Control?&quot;</span>,
<a name="l00192"></a>00192                                                        0, QApplication::UnicodeUTF8);
<a name="l00193"></a>00193     <a class="code" href="classSystemTracking_1_1CMainControl.html#ae3a35ef6c557bc66707db334a50a5280" title="slot for confirm module configuration.">ConfirmModuleConfiguration</a>(Text);
<a name="l00194"></a>00194 }
<a name="l00195"></a>00195 
<a name="l00196"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#a40d476648b94f3b5cc8892cbfe130f61">00196</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#a40d476648b94f3b5cc8892cbfe130f61" title="call slot when current tab changed.">CMainControl::CurrentTabChanged</a>(<span class="keywordtype">int</span> Index)
<a name="l00197"></a>00197 {
<a name="l00198"></a>00198     <span class="keywordflow">if</span> (Index != 0) {
<a name="l00199"></a>00199         <a class="code" href="classSystemTracking_1_1CMainControl.html#ae3a35ef6c557bc66707db334a50a5280" title="slot for confirm module configuration.">ConfirmModuleConfiguration</a>();
<a name="l00200"></a>00200     }
<a name="l00201"></a>00201 }
<a name="l00202"></a>00202 
<a name="l00203"></a><a class="code" href="classSystemTracking_1_1CMainControl.html#ae3a35ef6c557bc66707db334a50a5280">00203</a> <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#ae3a35ef6c557bc66707db334a50a5280" title="slot for confirm module configuration.">CMainControl::ConfirmModuleConfiguration</a>()
<a name="l00204"></a>00204 {
<a name="l00205"></a>00205     QString Text = QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00206"></a>00206                                            <span class="stringliteral">&quot;Main Control Module has been modified. Do you want to finalize the configuration?&quot;</span>,
<a name="l00207"></a>00207                                            0, QApplication::UnicodeUTF8);
<a name="l00208"></a>00208 
<a name="l00209"></a>00209     <span class="keywordflow">if</span> (mp_Ui-&gt;finalizeConfigBtn-&gt;isEnabled()) {
<a name="l00210"></a>00210         <a class="code" href="classSystemTracking_1_1CMainControl.html#ae3a35ef6c557bc66707db334a50a5280" title="slot for confirm module configuration.">ConfirmModuleConfiguration</a>(Text);
<a name="l00211"></a>00211     }
<a name="l00212"></a>00212 }
<a name="l00213"></a>00213 
<a name="l00214"></a>00214 <span class="keywordtype">void</span> <a class="code" href="classSystemTracking_1_1CMainControl.html#ae3a35ef6c557bc66707db334a50a5280" title="slot for confirm module configuration.">CMainControl::ConfirmModuleConfiguration</a>(QString&amp; Text)
<a name="l00215"></a>00215 {
<a name="l00216"></a>00216     mp_MessageDlg-&gt;SetButtonText(1, QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00217"></a>00217                                                             <span class="stringliteral">&quot;Ok&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00218"></a>00218     mp_MessageDlg-&gt;SetButtonText(3,QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00219"></a>00219                                                            <span class="stringliteral">&quot;Cancel&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00220"></a>00220     mp_MessageDlg-&gt;HideCenterButton();
<a name="l00221"></a>00221 
<a name="l00222"></a>00222     mp_MessageDlg-&gt;SetText(Text);
<a name="l00223"></a>00223     mp_MessageDlg-&gt;SetIcon(QMessageBox::Warning);
<a name="l00224"></a>00224     mp_MessageDlg-&gt;show();
<a name="l00225"></a>00225 
<a name="l00226"></a>00226     <span class="keywordtype">int</span> Result = mp_MessageDlg-&gt;exec();
<a name="l00227"></a>00227 
<a name="l00228"></a>00228     ResetMessageBox();
<a name="l00229"></a>00229     <span class="keywordflow">if</span> (Result)
<a name="l00230"></a>00230     {
<a name="l00231"></a>00231         <span class="keywordflow">if</span>(mp_DateConnector-&gt;UpdateInstrumentHistory(*mp_ModuleList))
<a name="l00232"></a>00232         {
<a name="l00233"></a>00233             emit <a class="code" href="classSystemTracking_1_1CMainControl.html#a170db7621bebcfbad548d4fda46bdc77" title="signal for module list changed.">ModuleListChanged</a>();
<a name="l00234"></a>00234             mp_MessageDlg-&gt;SetButtonText(1, QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00235"></a>00235                                                                     <span class="stringliteral">&quot;Ok&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00236"></a>00236             mp_MessageDlg-&gt;HideButtons();
<a name="l00237"></a>00237             mp_MessageDlg-&gt;SetText(QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00238"></a>00238                                           <span class="stringliteral">&quot;Configuration file updated successfully.&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00239"></a>00239             mp_MessageDlg-&gt;SetIcon(QMessageBox::Information);
<a name="l00240"></a>00240             mp_MessageDlg-&gt;show();
<a name="l00241"></a>00241         }
<a name="l00242"></a>00242         <span class="keywordflow">else</span>
<a name="l00243"></a>00243         {
<a name="l00244"></a>00244             mp_MessageDlg-&gt;SetButtonText(1, QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00245"></a>00245                                                                     <span class="stringliteral">&quot;Ok&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00246"></a>00246             mp_MessageDlg-&gt;HideButtons();
<a name="l00247"></a>00247             mp_MessageDlg-&gt;SetText(QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00248"></a>00248                                           <span class="stringliteral">&quot;Configuration file updation failed.&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00249"></a>00249             mp_MessageDlg-&gt;SetIcon(QMessageBox::Information);
<a name="l00250"></a>00250             mp_MessageDlg-&gt;show();
<a name="l00251"></a>00251         }
<a name="l00252"></a>00252     }
<a name="l00253"></a>00253     <span class="keywordflow">else</span>
<a name="l00254"></a>00254     {
<a name="l00255"></a>00255         <span class="comment">//Global::EventObject::Instance().RaiseEvent(EVENT_GUI_ADDMODIFY_SAVE_AND_OVERWRITE_CONFIGURATION_FAILURE);</span>
<a name="l00256"></a>00256         mp_MessageDlg-&gt;SetButtonText(1, QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00257"></a>00257                                                                 <span class="stringliteral">&quot;Ok&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00258"></a>00258         mp_MessageDlg-&gt;HideButtons();
<a name="l00259"></a>00259         mp_MessageDlg-&gt;SetText(QApplication::translate(<span class="stringliteral">&quot;SystemTracking::CMainControl&quot;</span>,
<a name="l00260"></a>00260                                              <span class="stringliteral">&quot;Finalize Configuration Cancelled.&quot;</span>, 0, QApplication::UnicodeUTF8));
<a name="l00261"></a>00261         mp_MessageDlg-&gt;SetIcon(QMessageBox::Warning);
<a name="l00262"></a>00262         mp_MessageDlg-&gt;show();
<a name="l00263"></a>00263     }
<a name="l00264"></a>00264     mp_Ui-&gt;finalizeConfigBtn-&gt;setEnabled(<span class="keyword">false</span>);
<a name="l00265"></a>00265 }
<a name="l00266"></a>00266 
<a name="l00267"></a>00267 <span class="keywordtype">void</span> CMainControl::AutoDetect(ServiceDataManager::CSubModule &amp;SubModule)
<a name="l00268"></a>00268 {
<a name="l00269"></a>00269     qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::AutoDetect : &quot;</span> &lt;&lt; SubModule.GetSubModuleName();
<a name="l00270"></a>00270 
<a name="l00271"></a>00271     Service::ModuleTestCaseID Id = Service::FIRMWARE_GET_SLAVE_INFO;
<a name="l00272"></a>00272     QString TestCaseName = DataManager::CTestCaseGuide::Instance().GetTestCaseName(Id);
<a name="l00273"></a>00273     DataManager::CTestCase* p_TestCase = DataManager::CTestCaseFactory::Instance().GetTestCase(TestCaseName);
<a name="l00274"></a>00274 
<a name="l00275"></a>00275     QString SubModuleName = SubModule.GetSubModuleName();
<a name="l00276"></a>00276     QString SlaveType = SubModuleName.remove(<span class="stringliteral">&quot;ASB&quot;</span>);
<a name="l00277"></a>00277     p_TestCase-&gt;SetParameter(<span class="stringliteral">&quot;SlaveType&quot;</span>, SlaveType);
<a name="l00278"></a>00278 
<a name="l00279"></a>00279     emit <a class="code" href="classSystemTracking_1_1CMainControl.html#a9e6df472eeaeb35d0d202b91945fbe92" title="Signal emitted to perform manufacturing tests.">PerformManufacturingTest</a>(Id);
<a name="l00280"></a>00280 }
<a name="l00281"></a>00281 
<a name="l00282"></a>00282 <span class="keywordtype">void</span> CMainControl::ModifySubModule(<span class="keyword">const</span> QString &amp;ModuleName,
<a name="l00283"></a>00283                                    <span class="keyword">const</span> QString &amp;SubModuleName)
<a name="l00284"></a>00284 {
<a name="l00285"></a>00285     <span class="keywordflow">if</span> (!mp_ModuleList) {
<a name="l00286"></a>00286         mp_ModuleList = <span class="keyword">new</span> ServiceDataManager::CModuleDataList;
<a name="l00287"></a>00287         ServiceDataManager::CModuleDataList* ModuleList = mp_DateConnector-&gt;GetModuleListContainer();
<a name="l00288"></a>00288         <span class="keywordflow">if</span> (!ModuleList) {
<a name="l00289"></a>00289             qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::UpdateSubModule(): Invalid module list!&quot;</span>;
<a name="l00290"></a>00290             <span class="keywordflow">return</span>;
<a name="l00291"></a>00291         }
<a name="l00292"></a>00292 
<a name="l00293"></a>00293         *mp_ModuleList = *ModuleList;
<a name="l00294"></a>00294     }
<a name="l00295"></a>00295 
<a name="l00296"></a>00296     ServiceDataManager::CModule *pModule = mp_ModuleList-&gt;GetModule(ModuleName);
<a name="l00297"></a>00297     <span class="keywordflow">if</span> (0 == pModule)
<a name="l00298"></a>00298     {
<a name="l00299"></a>00299         qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::ModifySubModule(): Invalid module : &quot;</span>
<a name="l00300"></a>00300                  &lt;&lt; ModuleName;
<a name="l00301"></a>00301         <span class="keywordflow">return</span>;
<a name="l00302"></a>00302     }
<a name="l00303"></a>00303 
<a name="l00304"></a>00304     ServiceDataManager::CSubModule *pSubModule =
<a name="l00305"></a>00305             pModule-&gt;GetSubModuleInfo(SubModuleName);
<a name="l00306"></a>00306     <span class="keywordflow">if</span> (0 == pSubModule)
<a name="l00307"></a>00307     {
<a name="l00308"></a>00308         qDebug() &lt;&lt; <span class="stringliteral">&quot;CMainControl::ModifySubModule(): Invalid sub module : &quot;</span>
<a name="l00309"></a>00309                  &lt;&lt; SubModuleName;
<a name="l00310"></a>00310         <span class="keywordflow">return</span>;
<a name="l00311"></a>00311     }
<a name="l00312"></a>00312 
<a name="l00313"></a>00313     ServiceDataManager::CSubModule SubModule = *pSubModule;
<a name="l00314"></a>00314 
<a name="l00315"></a>00315     CDlgModifySubModule *dlg = <span class="keyword">new</span> CDlgModifySubModule(SubModule, SubModuleName.contains(<span class="stringliteral">&quot;ASB&quot;</span>));
<a name="l00316"></a>00316 
<a name="l00317"></a>00317     dlg-&gt;UpdateGUI();
<a name="l00318"></a>00318 
<a name="l00319"></a>00319     (void)connect(dlg,
<a name="l00320"></a>00320                   SIGNAL(AutoDetect(ServiceDataManager::CSubModule&amp;)),
<a name="l00321"></a>00321                   <span class="keyword">this</span>,
<a name="l00322"></a>00322                   SLOT(AutoDetect(ServiceDataManager::CSubModule&amp;)) );
<a name="l00323"></a>00323 
<a name="l00324"></a>00324     (void)connect(dlg,
<a name="l00325"></a>00325                   SIGNAL(<a class="code" href="classSystemTracking_1_1CMainControl.html#a49a9bf2719340b4d072118f4f6925c33" title="slot for update sub module.">UpdateSubModule</a>(ServiceDataManager::CSubModule&amp;)),
<a name="l00326"></a>00326                   <span class="keyword">this</span>,
<a name="l00327"></a>00327                   SLOT(<a class="code" href="classSystemTracking_1_1CMainControl.html#a49a9bf2719340b4d072118f4f6925c33" title="slot for update sub module.">UpdateSubModule</a>(ServiceDataManager::CSubModule&amp;)) );
<a name="l00328"></a>00328 
<a name="l00329"></a>00329     (void)dlg-&gt;exec();
<a name="l00330"></a>00330 
<a name="l00331"></a>00331     <span class="keyword">delete</span> dlg;
<a name="l00332"></a>00332 }
<a name="l00333"></a>00333 
<a name="l00334"></a>00334 <span class="keywordtype">void</span> CMainControl::ResetMessageBox()
<a name="l00335"></a>00335 {
<a name="l00336"></a>00336     <span class="keywordflow">if</span> (mp_MessageDlg) {
<a name="l00337"></a>00337         <span class="keyword">delete</span> mp_MessageDlg;
<a name="l00338"></a>00338         mp_MessageDlg = <span class="keyword">new</span> MainMenu::CMessageDlg(<span class="keyword">this</span>);
<a name="l00339"></a>00339     }
<a name="l00340"></a>00340 }
<a name="l00341"></a>00341 
<a name="l00342"></a>00342 } <span class="comment">// namespace SystemTracking</span>
</pre></div></div><!-- contents -->
</div>
<!-- window showing the filter options -->
<div id="MSearchSelectWindow"
     onmouseover="return searchBox.OnSearchSelectShow()"
     onmouseout="return searchBox.OnSearchSelectHide()"
     onkeydown="return searchBox.OnSearchSelectKey(event)">
<a class="SelectItem" href="javascript:void(0)" onclick="searchBox.OnSelectItem(0)"><span class="SelectionMark">&#160;</span>All</a><a class="SelectItem" href="javascript:void(0)" onclick="searchBox.OnSelectItem(1)"><span class="SelectionMark">&#160;</span>Classes</a><a class="SelectItem" href="javascript:void(0)" onclick="searchBox.OnSelectItem(2)"><span class="SelectionMark">&#160;</span>Files</a><a class="SelectItem" href="javascript:void(0)" onclick="searchBox.OnSelectItem(3)"><span class="SelectionMark">&#160;</span>Functions</a></div>

<!-- iframe showing the search results (closed by default) -->
<div id="MSearchResultsWindow">
<iframe src="javascript:void(0)" frameborder="0" 
        name="MSearchResults" id="MSearchResults">
</iframe>
</div>

  <div id="nav-path" class="navpath">
    <ul>
      <li class="navelem"><a class="el" href="MainControl_8cpp.html">MainControl.cpp</a>      </li>

    <li class="footer">Generated on Thu Jul 24 2014 13:25:57 for Himalaya Service Master by
    <a href="http://www.doxygen.org/index.html">
    <img class="footer" src="doxygen.png" alt="doxygen"/></a> 1.7.6.1 </li>
   </ul>
 </div>


</body>
</html>
